# Miniproyecto 3: Clasificación de animales

## Entrega 2
En esta entrega, nos enfocaremos en aprender a utilizar nuestros dos primeros descriptores y en implementar nuestro primer clasificador. En particular, trabajaremos con el clasificador de K Vecinos Más Cercanos (K-NN). Los descriptores que utilizaremos en esta ocasión son: Descriptores de color y Descriptores de textura.

Es importante entender que no todas las formas de representar las imágenes son igualmente efectivas para la tarea de clasificación. El objetivo es encontrar la representación más adecuada de las imágenes, aquella que permita una clasificación precisa y eficiente. Para lograr esto, debemos probar y evaluar distintos tipos de descriptores.

En cuanto a los clasificadores, el clasificador de K Vecinos Más Cercanos (K-NN) es un algoritmo intuitivo y respaldado teóricamente, pero presenta algunas limitaciones. Una de las principales desventajas es que es muy sensible al tipo de distancia que utilicemos entre los puntos (o vectores de características) y también a la presencia de outliers en los datos. Además, K-NN es computacionalmente costoso, especialmente cuando trabajamos con grandes conjuntos de datos, ya que requiere comparar cada punto con todos los demás durante el proceso de clasificación.

### Preliminares
Copie las funciones necesarias de la entrega pasada en las siguientes celdas.

```python
# Importar librerias importantes
import os
import random
import pickle
import numpy as np
from tqdm import tqdm
from glob import glob
from matplotlib import pyplot as plt

from scipy.signal import correlate2d
from scipy.io import loadmat

import cv2

from sklearn.metrics import precision_recall_fscore_support
from sklearn.cluster import KMeans
from sklearn.neighbors import KNeighborsClassifier
```

<!-- grade_id: cell-072f4374cd451291, points: 0, grade: false, solution: true -->
```python
group = 13
```

```python
def determine_classes(group: int) -> list[str]:
    """Función que determina las clases a utilizar en el proyecto.

    Args:
        group (int): Número de grupo.

    Returns:
        list: Lista con los nombres de las clases a utilizar.
    """
    random.seed(group)
    categories = glob(os.path.join("dataset", "train", "*"))
    stats = [(os.path.basename(category), len(os.listdir(category))) for category in categories]
    clusters = {}
    for category, num_images in stats:
        if num_images in clusters:
            clusters[num_images].append(category)
        else:
            clusters[num_images] = [category]
    categories = []
    for cluster_categories in clusters.values():
        categories += [random.choice(cluster_categories)]
    return categories

assert group is not None, "Por favor, asigna un número de grupo a la variable 'group'."
assert group > 0, "El número de grupo debe ser mayor que 0."
assert group <= 28, "El número de grupo debe ser menor o igual que 28."

categories = determine_classes(group)
print(f"Las classes que usarán en este Mini-proyecto serán:\n{", ".join(categories)}.")
```

<!-- grade_id: cell-6e846061888ec7bb, points: 0, grade: false, solution: true -->
```python
dict_labels = {} # Cambiar por un diccionario con las etiquetas de las clases 

# YOUR CODE HERE
raise NotImplementedError()

def get_data(subset: str) -> tuple[list[str], list[int]]:
  """Función que obtiene las rutas de las imágenes y sus etiquetas.

  Args:
      subset (str): Nombre del subconjunto de datos a utilizar.

  Returns:
      tuple[list, list]: Rutas de las imágenes y sus etiquetas.
  """
  paths = []
  labels = []
  
  # YOUR CODE HERE
  raise NotImplementedError()
  return paths, labels

train_paths, train_labels = get_data("train")
valid_paths, valid_labels = get_data("valid")
test_paths, test_labels = get_data("test")
```

<!-- grade_id: cell-86830ced769cd3b7, points: 0.2, grade: true, solution: false -->
```python
assert len(train_paths) > 0, "No se ha cargado ningún dato de entrenamiento"
assert len(train_paths) == 5163, "El número de datos de entrenamiento no es correcto"
assert len(valid_paths) > 0, "No se ha cargado ningún dato de validación"
assert len(valid_paths) == 1723, "El número de datos de validación no es correcto"
assert len(test_paths) > 0, "No se ha cargado ningún dato de test"
assert len(test_paths) == 1724, "El número de datos de test no es correcto"

assert type(train_paths[0]) == str, "Los datos de entrenamiento no tienen la forma correcta"
assert type(valid_paths[0]) == str, "Los datos de validación no tienen la forma correcta"
assert type(test_paths[0]) == str, "Los datos de test no tienen la forma correcta"

assert type(train_labels[0]) == int, "Las etiquetas de entrenamiento no tienen la forma correcta"
assert type(valid_labels[0]) == int, "Las etiquetas de validación no tienen la forma correcta"
assert type(test_labels[0]) == int, "Las etiquetas de test no tienen la forma correcta"

assert len(dict_labels) == len(categories), "El diccionario de labels no tiene la longitud correcta"

categories_sorted = sorted(categories)
assert list(dict_labels.keys()) == categories_sorted, "Las llaves del diccionario de labels no son correctas"
assert list(dict_labels.values()) == list(range(len(categories))), "Los valores del diccionario de labels no son correctos"

assert min([dict_labels[train_paths[i].split(os.sep)[-2]] == train_labels[i] for i in range(len(train_paths))]), "Las etiquetas de entrenamiento no coinciden"
assert min([dict_labels[valid_paths[i].split(os.sep)[-2]] == valid_labels[i] for i in range(len(valid_paths))]), "Las etiquetas de validación no coinciden"
assert min([dict_labels[test_paths[i].split(os.sep)[-2]] == test_labels[i] for i in range(len(test_paths))]), "Las etiquetas de test no coinciden"
```

<!-- grade_id: cell-3c336429a1ca9159, points: 0, grade: false, solution: true -->
```python
def processing(Im: np.ndarray, img_size: tuple[int] | None = None, color_space: str | None = None) -> np.ndarray:
    """Función que preprocesa una lista de imágenes.

    Args:
        Im_list (list[np.ndarray]): Lista con las imágenes a preprocesar.
        img_size (tuple[int] | None): Tamaño de la imagen.
        color_space (str | None): Espacio de color a utilizar. Puede ser "rgb", "gray", "hsv" o "lab".

    Returns:
        list[np.ndarray]: Lista con las imágenes preprocesadas.
    """
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return Im
```

<!-- grade_id: cell-58bd2da03da76817, points: 0.2, grade: true, solution: false -->
```python
Im_list = [cv2.imread(path) for path in train_paths[:5]]
img_size = (224, 224)
color_space = "rgb"
Im_proc = [processing(Im, img_size, color_space) for Im in Im_list]

assert len(Im_proc) == len(Im_list), "El número de imágenes preprocesadas no es correcto"
assert type(Im_proc) == list, "Las imágenes preprocesadas no están en una lista"
assert type(Im_proc[0]) == np.ndarray, "Las imágenes preprocesadas no son arreglos de numpy"
assert np.array(Im_proc).shape[1:] == (224, 224, 3), "Las imágenes preprocesadas no tienen el tamaño correcto"
assert np.array_equal(cv2.cvtColor(cv2.resize(Im_list[0], (224, 224)), cv2.COLOR_BGR2RGB), Im_proc[0]), "Las imágenes preprocesadas no son correctas"

color_space = "gray"
Im_proc = [processing(Im, img_size, color_space) for Im in Im_list]
assert np.array(Im_proc).shape[1:] == (224, 224), "Las imágenes preprocesadas no tienen el tamaño correcto"
assert np.array_equal(cv2.cvtColor(cv2.resize(Im_list[0], (224, 224)), cv2.COLOR_BGR2GRAY), Im_proc[0]), "Las imágenes preprocesadas no son correctas"

color_space = "hsv"
Im_proc = [processing(Im, img_size, color_space) for Im in Im_list]
assert np.array(Im_proc).shape[1:] == (224, 224, 3), "Las imágenes preprocesadas no tienen el tamaño correcto"
assert np.array_equal(cv2.cvtColor(cv2.resize(Im_list[0], (224, 224)), cv2.COLOR_BGR2HSV), Im_proc[0]), "Las imágenes preprocesadas no son correctas"

color_space = "lab"
Im_proc = [processing(Im, img_size, color_space) for Im in Im_list]
assert np.array(Im_proc).shape[1:] == (224, 224, 3), "Las imágenes preprocesadas no tienen el tamaño correcto"
assert np.array_equal(cv2.cvtColor(cv2.resize(Im_list[0], (224, 224)), cv2.COLOR_BGR2LAB), Im_proc[0]), "Las imágenes preprocesadas no son correctas"

Im_proc = [processing(Im, (256, 256)) for Im in Im_list]
assert np.array(Im_proc).shape[1:] == (256, 256, 3), "Las imágenes preprocesadas no tienen el tamaño correcto"
assert np.array_equal(cv2.resize(Im_list[0], (256, 256)), Im_proc[0]), "Las imágenes preprocesadas no son correctas"

Im_proc = [processing(Im, None, "lab") for Im in Im_list]
assert np.array_equal(cv2.cvtColor(Im_list[0], cv2.COLOR_BGR2LAB), Im_proc[0]), "Las imágenes preprocesadas no son correctas"
```

### Parte 1: Histogramas
La representación de una imagen mediante su color puede hacerse de diversas maneras. La forma más común es usar histogramas de color, ya sea concatenados o conjuntos, para describir la distribución de colores en una imagen. Estos métodos estadísticos ya los hemos utilizado en el pasado, y para nuestro primer enfoque de clasificación usaremos un enfoque basado en histogramas de color. Sin embargo, es importante tener en cuenta que este método tiene limitaciones, ya que en ocasiones los colores de los animales o las escenas entre diferentes categorías pueden ser similares.

#### Parte 1.1: Color
Instrucciones:

1. Implementa una función que reciba una imagen de entrada y una cantidad específica de bins y retorne su histograma de color. El histograma puede ser:

    - Conjunto: Un histograma donde se consideren todas las combinaciones posibles de los canales de color (es decir, un histograma en un espacio tridimensional en el caso de RGB o Lab).
    - Concatenado: Un histograma independiente por cada canal de color, concatenando los histogramas de cada canal en un solo vector.

2. Retorno: El histograma debe estar normalizado, es decir, los valores deben estar escalados para que la suma total sea 1.

**Nota:** Pueden tomar la implementación del tutorial de la semana 11.

<!-- grade_id: cell-5c3e91532be1a9c7, points: 0, grade: false, solution: true -->
```python
def color_hist(Im: np.ndarray, type_h: str, bins: int) -> np.ndarray:
    """Función que calcula el histograma de color conjunto o concatenado de una imagen

    Args:
        Im (numpy.ndarray): Imagen a la que calcular el histograma
        type_h (str): Tipo de histograma a calcular, puede ser "concat" o "joint"
        bins (int): Cantidad de bins que usar en el histograma. En el caso de concat el número de bins finale será 3*bins
        En el caso de Joint será bins^3
        
    Returns:
        numpy.ndarray: Histograma de color normalizado
    """
    h = "" # Histograma de color normalizado
    # YOUR CODE HERE
    raise NotImplementedError()
    return h
```

<!-- grade_id: cell-e29adae74b2a314c, points: 0.2, grade: true, solution: false -->
```python
img=processing(cv2.imread(os.path.join(".", "dataset", "train", "Canis latrans", "f494fab00ebb2d87c8d3bf5215560d54.jpg")),  (224, 224), 'rgb')
concat_h=color_hist(img,'concat',10)
joint_h=color_hist(img,'joint',10)


fig,ax=plt.subplots(1,3,figsize=(15,5))
ax[0].hist(list(range(len(concat_h))), list(range(len(concat_h - 1))), weights = concat_h)
ax[0].set_title('Histograma concatenado')
ax[1].hist(list(range(len(joint_h))), list(range(len(joint_h - 1))), weights = joint_h)
ax[1].set_title('Histograma conjunto')
ax[2].imshow(img)
ax[2].axis('off')
ax[2].set_title('Imagen original')

plt.tight_layout()

assert np.isclose(np.sum(concat_h),1,0.1), 'Su histograma concatenado no esta normalizado'
assert np.isclose(np.sum(joint_h),1,0.1), 'Su histograma conjunto no esta normalizado'
assert len(concat_h.shape)==1 and len(joint_h.shape)==1, 'Su histograma debe tener una unica dimension'
assert np.isclose(len(concat_h)/3,len(joint_h)**(1/3),0.01),' Sus histogramas no tienen la cantidad de bins que deberían'
assert np.isclose(np.std(concat_h),0.024,0.05), 'Su histograma de color concatenado es erroneo'
assert np.isclose(np.std(joint_h),0.0063,0.05), 'Su histograma de color conjunto es erroneo'
```

#### Parte 1.2: Entrenamiento
Ahora que contamos con nuestra función para calcular los histogramas (también conocidos como descriptores) de una imagen, podemos utilizarlos para realizar la clasificación de las imágenes. Para ello, necesitamos dos componentes clave: los descriptores y un clasificador.

Proceso de extracción de descriptores:

1. Cálculo de histogramas:

    - Para cada imagen seleccionada, calcula tanto el histograma conjunto como el histograma concatenado.
    - Guarda las etiquetas (labels) correspondientes a cada imagen en otra variable. Asegúrate de que cada imagen tenga su histograma y etiqueta asociados.
    
    Nota: La cantidad de bins para los histogramas deben ser seleccionados según tu criterio. Es posible que al usar un número reducido de imágenes de entrenamiento no observes una gran diferencia entre los diferentes espacios de color. Recuerden utilizar su función de procesamiento para estandarizar las imágenes.

2. Imágenes de validación:

    - Aplica el mismo procedimiento de cálculo de histogramas a todas las imágenes del conjunto de validación (valid). Asegúrate de calcular los histogramas y guardar las etiquetas correspondientes de la misma forma que hiciste con las imágenes de entrenamiento.

Puntos importantes:

- Recuerda que la cantidad de bins afecta el desempeño. Puedes experimentar con diferentes configuraciones para observar cómo influyen en la clasificación.

<!-- grade_id: cell-61d90fb11ea7b02c, points: 0, grade: false, solution: true -->
```python
train_images_num='' #cantidad de imagens de entrenamiento
train_joint_color_desc=[] #Descriptor de imagenes con histogramas conjuntos de color
train_concat_color_desc=[] #Descriptor de imagenes con histogramas concatenados de color

valid_joint_color_desc=[] #Descriptor de imagenes con histogramas conjuntos de color
valid_concat_color_desc=[] #Descriptor de imagenes con histogramas concatenados de color

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b48fd731d21fab11, points: 0.4, grade: true, solution: false -->
```python
assert len(train_joint_color_desc) == 5163, 'La cantidad de descriptores no es igual a la cantidad de imágenes totales'
assert len(train_concat_color_desc) == 5163, 'La cantidad de descriptores no es igual a la cantidad de imágenes totales'
assert len(valid_joint_color_desc) == 1723, 'La cantidad de descriptores no es igual a la cantidad de imágenes totales'
assert len(valid_concat_color_desc) == 1723, 'La cantidad de descriptores no es igual a la cantidad de imágenes totales'

assert len(train_joint_color_desc[0])<= 7056, 'Su descriptor unimodal de color tiene demasiadas dimensiones'
assert len(train_concat_color_desc[0])<= 7056, 'Su descriptor unimodal de color tiene demasiadas dimensiones'
assert len(train_joint_color_desc[0]) == len(train_joint_color_desc[0]), 'Los descriptores conjuntos tienen diferentes dimensiones'
assert len(train_concat_color_desc[0]) == len(valid_concat_color_desc[0]), 'Los descriptores concatenados tienen diferentes dimensiones'
```

Ahora que hemos calculado los descriptores para varias imágenes y utilizaremos el clasificador [KNeighborsClassifier](https://scikit-learn.org/stable/modules/generated/sklearn.neighbors.KNeighborsClassifier.html) para realizar el entrenamiento con los descriptores y etiquetas del conjunto de entrenamiento (train). El clasificador de K vecino más cercano debería ser un concepto familiar para ustedes a este punto. Sin embargo, repasemos brevemente su funcionamiento. El clasificador de vecino más cercano compara un nuevo dato (descriptor) con todos los datos de entrenamiento y le asigna la etiqueta de la clase predominante dentro de los K vecinos más cercanos. Luego, utilizaremos los descriptores del conjunto de validación (valid) para predecir las etiquetas correspondientes. Finalmente, evaluaremos qué tan buenos son nuestros resultados.

**Proceso:**
1. Construcción de clasificadores:

    - Crea un objeto del tipo KNeighborsClassifier para cada uno de los dos tipos de descriptores calculados (histograma conjunto e histograma concatenado).

2. Entrenamiento:

    - Utiliza el método .fit de cada modelo para entrenar los clasificadores con los descriptores y las etiquetas del conjunto de entrenamiento.

3. Predicción:

    - Aplica el método .predict a los descriptores del conjunto de validación para cada uno de los modelos (histograma conjunto y concatenado) y obtén las etiquetas predichas.

<!-- grade_id: cell-e112bea36409a92e, points: 0, grade: false, solution: true -->
```python
NNClassifier_concatcolor='' # Clasificador con histogramas concatenados de color
NNClassifier_jointcolor='' # Clasificador con histogramas conjuntos de color

predict_NNClassifier_concatcolor = '' # Predicciones del clasificador con histogramas concatenados de color
predict_NNClassifier_jointcolor = '' # Predicciones del clasificador con histogramas conjuntos de color

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-59e873620d709a70, points: 0.5, grade: true, solution: false -->
```python
assert NNClassifier_concatcolor.n_samples_fit_ == 5163, 'No se entrenó el clasificador con histogramas concatenados de color'
assert NNClassifier_concatcolor.n_features_in_ > 0, 'No se entrenó el clasificador con histogramas concatenados de color'
assert NNClassifier_jointcolor.n_samples_fit_ == 5163, 'No se entrenó el clasificador con histogramas conjuntos de color'
assert NNClassifier_jointcolor.n_features_in_ > 0, 'No se entrenó el clasificador con histogramas conjuntos de color'
assert NNClassifier_concatcolor.classes_.shape[0] == len(categories), 'El clasificador con histogramas concatenados de color no tiene las clases correctas'
assert NNClassifier_jointcolor.classes_.shape[0] == len(categories), 'El clasificador con histogramas conjuntos de color no tiene las clases correctas'

assert len(predict_NNClassifier_concatcolor) == len(valid_concat_color_desc), 'El número de predicciones debe ser igual al número de descriptores de validación'
assert len(predict_NNClassifier_jointcolor) == len(valid_joint_color_desc), 'El número de predicciones debe ser igual al número de descriptores de validación'
```

Ahora que hemos entrenado nuestros modelos y obtenido las predicciones, evaluaremos su rendimiento utilizando las métricas de precisión, recall y f1-score.

Utilizaremos la función precision_recall_fscore_support de la librería sklearn.metrics para calcular estas métricas comparando las etiquetas reales con las predicciones.

**Instrucciones:**

1. Evaluación de los modelos:

    - Para cada modelo (histograma conjunto e histograma concatenado), aplica precision_recall_fscore_support para calcular las métricas de precisión, recall y f1-score, comparando las etiquetas verdaderas del conjunto de validación con las predicciones obtenidas.

2. Uso del parámetro average:

    - Si deseas obtener el promedio de las métricas, utiliza average="macro".
    - Si prefieres obtener las métricas individuales para cada categoría, utiliza average=None. Esto te devolverá las métricas por separado para cada categoría de animal.
    
Nota: Finalmente, compara los resultados obtenidos para ambos modelos y analiza cuál funciona mejor en función de las métricas calculadas.

<!-- grade_id: cell-bdb5c160e473668c, points: 0, grade: false, solution: true -->
```python
precision_concatcolor = '' #Precision del clasificador con histogramas concatenados de color
recall_concatcolor = '' #Recall del clasificador con histogramas concatenados de color
f1_score_concatcolor = '' #F1 del clasificador con histogramas concatenados de color

precision_jointcolor = '' #Precision del clasificador con histogramas conjuntos de color
recall_jointcolor = '' #Recall del clasificador con histogramas conjuntos de color
f1_score_jointcolor = '' #F1 del clasificador con histogramas conjuntos de color

# YOUR CODE HERE
raise NotImplementedError()

print ("{:<25} | {:<14} {:<14} {:<14}".format('Modelo','Precisión','Cobertura', 'F1'))
print ('----------------------------------------------------------------------')
print ("{:<25} | {:<14.3f} {:<14.3f} {:<14.3f}".format('Histograma concatenado',precision_concatcolor,recall_concatcolor,f1_score_concatcolor))
print ("{:<25} | {:<14.3f} {:<14.3f} {:<14.3f}".format('Histograma Conjunto',precision_jointcolor,recall_jointcolor,f1_score_jointcolor))
```

<!-- grade_id: cell-83330e0b5c805117, points: 0.4, grade: true, solution: false -->
```python
assert f1_score_concatcolor > 0.28, 'El modelo con histogramas concatenados de color debería tener una f-medida mayor'
assert f1_score_jointcolor > 0.28, 'El modelo con histogramas cnojuntos de color debería tener una f-medida mayor'
```

### Parte 2: Textura
En la sección anterior ya utilizamos el color como representación de nuestras imágenes. Dependiendo del problema, esta puede ser una representación válida, ya que algunos animales pueden ser fácilmente diferenciados por su color. Sin embargo, en muchos casos, animales de distintas categorías pueden compartir un color similar, o bien, animales de la misma categoría pueden tener colores muy diferentes.

Una alternativa interesante es utilizar un descriptor de textura para representar nuestras imágenes. La textura se refiere a patrones repetitivos dentro de las imágenes, como el pelaje de un animal o las escamas en la piel. Al capturar estos patrones, podemos encontrar una forma distinta y, en algunos casos, más efectiva de representar las imágenes para su clasificación.

En esta parte, aprenderemos cómo representar una imagen mediante su textura.

Primero implementaremos un código que selecciona aleatoriamente 25 imágenes del conjunto de entrenamiento. Esas 25 imágenes aleatorias las guardaras en una lista llamada Im_list. Además, para ser consistentes utilice la función `processing` para cambiar el tamaño de las imágenes a $256\times256$ y convertirla a escala de grises.

<!-- grade_id: cell-d02dbf763a64b5f3, points: 0, grade: false, solution: true -->
```python
Im_list = '' #lista de imagenes, 5 de cada categoría
# YOUR CODE HERE
raise NotImplementedError()
```

#### Parte 2.1: Respuesta a filtros

Para detectar los patrones repetitivos dentro de una imagen (como bordes, barras, puntos, etc.), utilizaremos la respuesta de la imagen a diferentes filtros. Se te proporcionará un banco de filtros que contiene varios tipos de patrones.

El archivo `.mat` que recibirás contiene una matriz de dimensiones 49x49x18, lo que significa que son 18 filtros de tamaño 49x49 cada uno.

Instrucciones:
1. Utiliza la función loadmat de scipy.io para leer el archivo .mat. Esta función devolverá un diccionario, y la llave "filterbank" te proporcionará la matriz con los filtros.

**Nota:** Revisen bien cómo utilizar loadmat para cargar el archivo y cómo acceder a la matriz de filtros. Esto será esencial para los próximos pasos del miniproyecto.



<!-- grade_id: cell-79695f8b99f4422e, points: 0, grade: false, solution: true -->
```python
filterbank='' #arreglo de filtros
# YOUR CODE HERE
raise NotImplementedError()
```

Tu tarea ahora es crear una función que reciba por parámetro una imagen y el filterbank (en forma de array). La función debe aplicar cada uno de los N filtros del filterbank a la imagen de entrada, donde N es la cantidad de filtros presentes en el filterbank.

El resultado final debe ser una imagen con N canales, donde cada canal corresponde a la respuesta de la imagen a uno de los filtros. El tamaño de cada canal debe ser el mismo que el de la imagen original.

Instrucciones:
1. Utiliza la función correlate2d para realizar el filtrado de la imagen con cada filtro.
2. Importante: Asegúrate de verificar si la imagen está en escala de grises y si es necesario convertirla, ya que la respuesta de los filtros está diseñada para imágenes en gris.
4. Al final, combina las respuestas de todos los filtros en una sola imagen con N canales.

**Consejo:** `correlate2d` es una función de scipy.signal que te permitirá realizar la correlación de la imagen con cada filtro. Usando el modo "same", te asegurarás de que la salida de cada filtro tenga las mismas dimensiones que la imagen original.

<!-- grade_id: cell-8142635bac721a0e, points: 0, grade: false, solution: true -->
```python
def filter_map(Im: np.ndarray, filterbank: np.ndarray) -> np.ndarray:
    """Funcion que aplica un banco de filtros a una imagen

    Args:
        Im (numpy.ndarray): Imagen de entrada RGB
        filterbank (numpy.ndarray): Banco de filtros

    Returns:
        numpy.ndarray: Resultado de aplicar el banco de filtros a la imagen
    """
    new_im = '' #Imagen con los filtros aplicados
    # YOUR CODE HERE
    raise NotImplementedError()
    return new_im
```

<!-- grade_id: cell-a180f2dd148098f3, points: 0.4, grade: true, solution: false -->
```python
img=processing(cv2.imread(os.path.join(".", "dataset", "train", "Canis latrans", "f494fab00ebb2d87c8d3bf5215560d54.jpg")), (256, 256), 'gray')
new_img=filter_map(img,filterbank)

x,y=img.shape
x2,y2,f=new_img.shape
a,b,f2=filterbank.shape

assert f2==f, 'La imagen resultante debe tener la misma cantidad de canales como filtros en el filterbank'
assert x==x2 and y==y2, 'Las dimensiones de la imagen original no deben cambiar'
assert np.sum(np.std(new_img,axis=2))!=0, 'Usted usó el mismo filtro para todas las dimensiones'

assert np.sum(new_img[:,:,0]-correlate2d(img,filterbank[:,:,0],'same'))==0, 'El primer canal de new_im no es el resultado del filtrado con el primer filtro'
assert np.sum(new_img[:,:,2]-correlate2d(img,filterbank[:,:,2],'same'))==0, 'El primer canal de new_im no es el resultado del filtrado con el tercer filtro'

try:
    img_1=cv2.resize(cv2.imread(os.path.join(".", "dataset", "train", "Canis latrans", "f494fab00ebb2d87c8d3bf5215560d54.jpg")), (256, 256))
    filter_map(img_1,filterbank)
except:
    assert False, 'Su función debe funcionar con imágenes a color'

fig,ax=plt.subplots(2,3,figsize=(15,10))
ax[0,0].set_title('Imagen original')
ax[0,0].imshow(img, cmap='gray')
ax[0,1].set_title('Imagen filtrada 1')
ax[0,1].imshow(new_img[:,:,0])
ax[0,2].set_title('Imagen filtrada 10')
ax[0,2].imshow(new_img[:,:,9])

ax[0,0].axis('off')
ax[0,1].axis('off')
ax[0,2].axis('off')

ax[1,1].imshow(filterbank[:,:,0])
ax[1,1].set_title('Filtro 1')
ax[1,2].imshow(filterbank[:,:,9])
ax[1,2].set_title('Filtro 10')

ax[1,0].axis('off')
ax[1,1].axis('off')
ax[1,2].axis('off')

plt.show()
```

#### Parte 2.2: Colección de pixeles

Ya hemos creado una función que calcula la respuesta de una imagen a los filtros de un banco de filtros. Sin embargo, por sí sola, esta información aún no es suficiente. Recordemos que las texturas son patrones repetitivos en una imagen, y nuestro objetivo es identificar esos patrones. Pero, como seres humanos, no siempre podemos definir exactamente qué constituye un patrón repetitivo ni cuántos patrones hay en una imagen. Por lo tanto, dejaremos que la computadora defina y agrupe estos patrones por nosotros.

Utilizaremos el algoritmo de KMeans para agrupar los patrones de textura presentes en la imagen. Para ello, usaremos la respuesta de cada píxel a los filtros como datos de entrada para el algoritmo.

**Proceso:**
1. Datos de entrenamiento: Los datos de entrada para KMeans serán las respuestas de cada píxel a los N filtros aplicados en la parte anterior. Es decir, cada píxel será representado como un vector de tamaño N, donde cada valor corresponde a la respuesta del píxel a un filtro específico.

2. KMeans: Aplicaremos el algoritmo de KMeans para agrupar los píxeles en diferentes clústeres de texturas. El número de clústeres que escojamos corresponderá a la cantidad de texturas que la computadora identificará en la imagen.

##### Parte 2.2.1: Pixel dataset

Crea una función que reciba una lista de imágenes junto con el filterbank. Esta función debe realizar los siguientes pasos:

1. Obtener el mapa de respuesta a los filtros para cada imagen. Utiliza la función previamente implementada que aplica los filtros a la imagen y genera un mapa de respuesta.

2. Extraer los píxeles del mapa de respuesta: De cada imagen, extrae los valores de cada píxel a través de los N canales (donde N es el número de filtros en el filterbank). Cada píxel se convertirá en un vector de N dimensiones.

3. Agregar los píxeles a una lista: Crea una lista donde cada elemento será uno de los píxeles extraídos con sus N canales.

**Ejemplo:**
- Si tu imagen es de tamaño 256x256 y aplicas 18 filtros, el mapa de respuesta tendrá un tamaño de 256x256x18, lo que significa que tendrás 65536 píxeles, cada uno representado por un vector de 18 dimensiones.
- Si tienes 5 imágenes de tamaño 256x256, al final tendrás 327,680 píxeles, cada uno con 18 dimensiones.

Nota: Asegúrate de que todas tus imágenes tengan un tamaño de 256x256 antes de aplicar los filtros. Puedes redimensionarlas si es necesario.

<!-- grade_id: cell-ef5e94c68a1e0507, points: 0, grade: false, solution: true -->
```python
def pixel_dataset(Im_list: list[np.ndarray], filterbank: np.ndarray) -> list[np.ndarray]:
    """Funcion que aplica un banco de filtros a una lista de imagenes y retorna un arreglo con los pixeles de las imagenes filtradas

    Args:
        Im_list (list[numpy.ndarray]): Lista de imágenes de entrada RGB
        filterbank (numpy.ndarray): Banco de filtros

    Returns:
        list[numpy.ndarray]: Lista con los pixeles de las imagenes filtradas
    """
    pixels = [] #Lista con los pixeles de las imagenes filtradas
    # YOUR CODE HERE
    raise NotImplementedError()
    return pixels
```

<!-- grade_id: cell-39fd5f9fcd0fecf5, points: 0.5, grade: true, solution: false -->
```python
Pixels=pixel_dataset(Im_list,filterbank)

assert len(Pixels)==256*256*25, f'Deben tener una cantidad de 1638400 datos, pues son 256x256 pixeles x 25 imagenes, usted tiene {len(Pixels)}'
assert len(Pixels[0])==18, 'Las dimensiones de cada pixel deben ser 18, uno por cada filtro'
assert 0 in np.sum(np.array(Pixels)-filter_map(Im_list[0],filterbank)[0,0,:],axis=1), 'Debe agregar todos los pixeles de la imagen a su lista, puede que haya agregado mal los pixeles'
```

##### Parte 2.2.2: Modelo Kmeans

Ya hemos obtenido la respuesta a los filtros para cada uno de los píxeles de nuestro subdataset. Ahora procederemos a entrenar un modelo KMeans, donde el número de clústeres K representará los textones del modelo, es decir, las unidades fundamentales de la textura y los patrones repetitivos que buscamos identificar.

**Instrucciones:**
1. Crea un modelo KMeans con K=25, donde K será el número de textones que deseas identificar. Este valor ha sido seleccionado para capturar una variedad suficiente de patrones de textura en las imágenes.

2. Entrena el modelo utilizando la lista de píxeles obtenida en la parte anterior.

3. Este modelo de KMeans actuará como un diccionario de textones, agrupando los píxeles en diferentes clústeres que representan patrones repetitivos de la textura.

<!-- grade_id: cell-6e9d650b5172b860, points: 0, grade: false, solution: true -->
```python
textones = 25 #cantidad de textones k
model1='' #modelo K means entrenado
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3dba8dba727ba486, points: 0.4, grade: true, solution: false -->
```python
assert len(model1.cluster_centers_) ==textones, 'Debe entrenar su modelo Kmeans con la cantidad de textones establecido'
assert model1.n_features_in_ == 18, 'Usted debe entrenar el modelo con el dataset de pixeles'
```

#### Parte 2.3: Predicción de texturas
Ya tenemos creado nuestro diccionario de textones con el modelo KMeans. Ahora, debes crear una función que realice la predicción de texturas en una imagen.

**Instrucciones:**

1. La función debe calcular la respuesta a los filtros de la imagen utilizando el filter_map que creaste previamente, es decir, la imagen resultante con N canales (uno por cada filtro aplicado).

2. Convierte el filter_map en una lista de píxeles, es decir, pasar de 256x256x25 a 65536x25.

3. Utiliza el modelo KMeans para predecir la etiqueta de cada píxel (asignando la etiqueta del clúster correspondiente).

4. Reconstruye la imagen original asignando a cada píxel la etiqueta predicha por KMeans, de manera que la imagen resultante tenga las mismas dimensiones que la original, pero con una sola etiqueta entera por píxel (representando el clúster al que pertenece).

5. Devuelve una imagen de 2 dimensiones, donde cada píxel tendrá la etiqueta del clúster al que fue asignado.

**Nota:** KMeans predice de manera más eficiente en paralelo. Es mucho más rápido predecir una lista completa de píxeles (por ejemplo, 1000 píxeles a la vez) que predecir cada píxel individualmente. Intenta convertir el filter_map en una lista de píxeles, predecirlos todos al mismo tiempo y luego reconstruir la imagen utilizando .reshape() para devolverla a sus dimensiones originales.

<!-- grade_id: cell-87ae7d474a2d66a2, points: 0, grade: false, solution: true -->
```python
def texture_map(Im: np.ndarray, filterbank: np.ndarray, texton_dict: KMeans) -> np.ndarray:
    """Funcion que mapea una imagen a textones

    Args:
        Im (numpy.ndarray): Imagen de entrada RGB
        filterbank (numpy.ndarray): Banco de filtros
        texton_dict (KMeans): Diccionario de textones

    Returns:
        numpy.ndarray: Imagen mapeada a textones
    """
    texture_im = '' # Imagen mapeada a textones
    # YOUR CODE HERE
    raise NotImplementedError()
    return texture_im
```

<!-- grade_id: cell-1fd7fd5fd31f8e3a, points: 0.5, grade: true, solution: false -->
```python
img=processing(cv2.imread(os.path.join(".", "dataset", "train", "Canis latrans", "f494fab00ebb2d87c8d3bf5215560d54.jpg")),  (256, 256), 'gray')
texture_img=texture_map(img,filterbank,model1)

x,y=img.shape

assert len(texture_img.shape)==2, 'Su imagen debería tener solo 2 dimensiones'
x2,y2=texture_img.shape
assert x==x2 and y==y2, 'Las dimensiones de su imagen no deberían cambiar'
assert model1.predict([filter_map(img,filterbank)[0,0,:]])==texture_img[0,0],'La predicción del pixel 0,0 es erroneo'
assert model1.predict([filter_map(img,filterbank)[24,24,:]])==texture_img[24,24],'La predicción del pixel 24,24 es erroneo'

fig,ax=plt.subplots(1,2,figsize=(10,10))
ax[0].imshow(cv2.cvtColor(img, cv2.COLOR_BGR2RGB))
ax[1].imshow(texture_img)

ax[0].axis('off')
ax[1].axis('off')
plt.show()
```

Por último, crearemos un histograma de textones para describir la imagen, basado en las etiquetas de textones generadas por KMeans.

**Instrucciones:**

La función debe seguir los siguientes pasos:

1. Utiliza la función texture_map que implementaste previamente para crear el mapa de texturas de la imagen. Este mapa es una versión de la imagen donde cada píxel está representado por la etiqueta del clúster al que pertenece, según KMeans.

2. Una vez tengas el mapa de texturas (una imagen de 2 dimensiones con las etiquetas de los clústeres), crea un histograma utilizando la cantidad de bins especificada y de rango (0, bins-1). Este histograma contará la frecuencia de cada textón en la imagen.

3. Normaliza el histograma, de manera que la suma total de los valores sea 1. Esto permitirá comparar el histograma entre diferentes imágenes independientemente de su tamaño.

**Nota:** El histograma debe reflejar la frecuencia de aparición de cada textón (clúster) en la imagen. Asegúrate de especificar correctamente la cantidad de bins para agrupar los textones.

<!-- grade_id: cell-7d69f1d58006f14b, points: 0, grade: false, solution: true -->
```python
def texture_histogram(Im: np.ndarray, filterbank: np.ndarray, texton_dict: KMeans, bins: int) -> np.ndarray:
    """Funcion que calcula el histograma de textones de una imagen

    Args:
        Im (numpy.ndarray): Imagen de entrada RGB
        filterbank (numpy.ndarray): Banco de filtros
        texton_dict (KMeans): Diccionario de textones
        bins (int): Número de bins del histograma

    Returns:
        numpy.ndarray: Histograma de textones normalizado
    """
    histogram = "" # Histograma de textones normalizado
    # YOUR CODE HERE
    raise NotImplementedError()
    return histogram
```

<!-- grade_id: cell-7b5ac3b3826ab556, points: 0.4, grade: true, solution: false -->
```python
histogram_text=texture_histogram(img,filterbank,model1,textones)
assert np.isclose(np.sum(histogram_text),1), 'El histograma no esta normalizado'
assert len(histogram_text)==textones, 'El histograma no tiene la cantidad de bins que se indicó'
```

#### Parte 2.4: Clasificación
En esta etapa, utilizaremos nuevamente el clasificador vecino más cercano pero con los descriptores de textura. Ahora que tenemos un método para la clasificación y un método para describir nuestras imágenes mediante vectores, es momento de combinar estas dos metodologías para crear nuestro modelo de clasificación. Obtendrás el descriptor de textura usando 25 textones (el modelo KMeans que creaste anteriormente) y 25 bins.

##### Parte 2.4.1: Datos

1. Obtener descriptores de textura:

    - Extrae el descriptor de 150 imágenes en el conjunto de entrenamiento (train).
    - Extrae el descriptor de las 50 (10 de cada clase) imágenes en el conjunto de validación (valid) junto con las labels correspondientes.

2. Guardar los datos:

    - Se recomienda usar la librería pickle para guardar los descriptores y etiquetas. Esto facilitará la carga de los datos en futuros pasos.
    
3. Ajuste de tamaño:

    - Recuerda hacer un resize con la función processing de las imágenes a un tamaño de 256x256 antes de extraer los descriptores.

<!-- grade_id: cell-4793545dc17379af, points: 0, grade: false, solution: true -->
```python
train_labels_list = [] #lista de etiquetas de entrenamiento
train_texture_desc=[] #vector de descriptores

valid_labels_list = [] #lista de etiquetas de validación
valid_texture_desc=[] #vector de descriptores
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-da4295ee888de9ca, points: 0.4, grade: true, solution: false -->
```python
assert len(valid_texture_desc[0]) == len(train_texture_desc[0]), 'Los descriptores no son de la misma dimensión'
assert len(valid_texture_desc) == len(valid_labels_list), 'Usted no cargó la cantidad correcta de etiquetas'
assert len(train_texture_desc) == len(train_labels_list), 'Usted no cargó la cantidad correcta de etiquetas'
```

Realiza la clasificación creando un objeto de la clase NN_classifier, entrena con .fit usando los descriptores de entrenamiento y predice los descriptores de validación con .predict.

<!-- grade_id: cell-812324d91fde98dd, points: 0, grade: false, solution: true -->
```python
predict_vector=[] #vector de etiquetas predichas
NNclassifier = '' #modelo de clasificador
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-04ddcac79e88eaf3, points: 0.3, grade: true, solution: false -->
```python
assert len(predict_vector)==len(valid_texture_desc), 'El vector de predicciones debe la misma cantidad que descriptores de validación'
assert NNclassifier.n_features_in_== 25, 'El modelo no tiene la cantidad de textones correcta'
assert NNclassifier.n_samples_fit_ == 150, 'El modelo no fue entrenado correctamente'
```

Finalmente, obtén las métricas de clasificación utilizando la función `precision_recall_fscore_support` de sklearn.

**Nota:** Usa average="macro" para obtener las métricas promediadas o average=None para obtenerlas para cada clase.

<!-- grade_id: cell-8184f4760721ac90, points: 0, grade: false, solution: true -->
```python
precision='' #Vector de precisión de cada categoría
recall='' #Vector de cobertura de cada categoría
f1_score='' #Vector de f medida de cada categoría
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2ab37fa675497c32, points: 0.2, grade: true, solution: false -->
```python
assert f1_score > 0.16, 'La f-medida no es la esperada'
```

### Textura forma eficiente
Cómo se pudieron dar cuenta los descriptores de textura pueden ser poco eficientes. El proceso se demora un tiempo considerable. Por lo tanto nosotros le proveeremos una función de `filter_map` que permite realizar todo el procedimiento de una manera mucho más rápida.

**Nota:** Para poder utilizar esta función de forma correcta tendrán que usar una librería llamada torch, por lo tanto le recomendamos usar Google Colab y configurar el entorno de ejecución para utilizar un GPU cómo acelerador por Hardware. El Asistente Graduado explicará esto de forma más detallada.

```python
import torch

def filter_map(Im: np.ndarray, filterbank: np.ndarray) -> torch.tensor:
    """Funcion que aplica un banco de filtros a una imagen

    Args:
        Im (numpy.ndarray): Imagen de entrada RGB

    Returns:
        numpy.ndarray: Resultado de aplicar el banco de filtros a la imagen
    """
    if len(Im.shape) == 3:
        Im = cv2.cvtColor(Im, cv2.COLOR_BGR2GRAY)
        
    device = torch.device("cuda" if torch.cuda.is_available() else "cpu")
    dtype = torch.float32

    Im = torch.tensor(Im, dtype=dtype, device=device)
    Im = Im[None, None]
    filterbank = torch.tensor(filterbank, dtype=dtype, device=device)
    filterbank = filterbank.permute(2, 0, 1)[:, None]
    convolution = torch.nn.functional.conv2d(Im, filterbank, padding="same")
    convolution = convolution.permute(0, 2, 3, 1)
    new_im = convolution.to("cpu").numpy()[0]

    return new_im
```

## Entregables
Los entregables para esta primera entrega son los siguientes:

- Jupyter notebook (.ipynb): El notebook debe estar completamente resuelto. Recuerden que la nota depende del correcto funcionamiento de su código.

- Informe (.pdf): Un informe que describa lo realizado en esta entrega, incluyendo visualizaciones, análisis y otros aspectos relevantes para demostrar el desarrollo del miniproyecto. La extensión máxima del informe será de 3 páginas. No es necesario que sea muy detallado ni que siga una estructura específica, solo se solicita para verificar que hayan cumplido correctamente con el desarrollo del miniproyecto.

- Archivo de texto (.txt): Un archivo de texto con las secciones de código del Jupyter notebook. Para más información sobre cómo convertir las secciones de código de su Jupyter notebook a un archivo de texto, pueden ver el video disponible en el siguiente [enlace](https://uniandes-edu-co.zoom.us/rec/share/QOxUUIw7Uz9DsnKPyftOXZgM4bx5d7KHMYK-dqvGnkyJ2GfqvT7NU3lhQc0NtNSU.jveWB2H34S4uLIoB).

**Nota:** En esta entrega no adelantaremos las secciones del artículo. Continuaremos con el artículo la próxima entrega cuando ya tengamos un método de base.

```python
from utils import *
converter("Entrega 2")
```